Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Assignment 4: List comprechension, generators, functional python

Imports

```python
from contextlib import redirect_stdout
from typing import Iterable, Any, List
from types import GeneratorType
from functools import reduce
import io

```

## 1. Comprehensions

Convert the following for loops into comprehensions:

```python
l = []
for i in range(-5, 10, 2):
    l.append(i-2)
l
```

Output:
```
[-7, -5, -3, -1, 1, 3, 5, 7]
```

<!-- grade_id: cell-fe81859f4057f7ff, points: 0, grade: false, solution: true -->
```python
l2 = [i-2 for i in range(-5, 10, 2)]
l2
```

Output:
```
[-7, -5, -3, -1, 1, 3, 5, 7]
```

```python
l = []
for i in range(100):
    if i % 10 == 4:
        l.append(i)
l
```

Output:
```
[4, 14, 24, 34, 44, 54, 64, 74, 84, 94]
```

<!-- grade_id: cell-adadaf2661704ec1, points: 0, grade: false, solution: true -->
```python
l2 = [i for i in range(100) if i % 10 == 4]
l2
```

Output:
```
[4, 14, 24, 34, 44, 54, 64, 74, 84, 94]
```

```python
l1 = [12, 1, 0, 13, -3, -4, 0, 2]
l2 = []

for e in l1:
    if e % 2 == 1:
        l2.append(e)
l2
```

<!-- grade_id: cell-01c992d67ca14968, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
l1 = [12, 1, 0, 13, -3, -4, 0, 2]
l2 = []

for e in l1:
    if e % 2 == 1:
        l2.append(True)
    else:
        l2.append(False)
l2
```

<!-- grade_id: cell-106d47d0b13dd3b1, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
l1 = [3, 5, 7, 11]
l2 = [2, 4, 6]

products = []

for x in l1:
    for y in l2:
        products.append(x*y)
products
```

Output:
```
[6, 12, 18, 10, 20, 30, 14, 28, 42, 22, 44, 66]
```

<!-- grade_id: cell-4ead65faba295143, points: 0, grade: false, solution: true -->
```python
prod2 = [x * y for x in l1 for y in l2]
prod2
```

Output:
```
[6, 12, 18, 10, 20, 30, 14, 28, 42, 22, 44, 66]
```

```python
l1 = [3, 5, 7, 11, 13, 17, 19]
l2 = [2, 4, 6, 8, 10]

products = []

for x in l1:
    for y in l2:
        if (x + y) % 3 == 0:
            products.append(x*y)
products
```

<!-- grade_id: cell-c8caea611b1907b3, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
fruits = ["apple", "plum", "pear", "avocado"]

mtx = []
for fruit in fruits:
    row = []
    for i, c in enumerate(fruit):
        row.append(c*(i+1))
    mtx.append(row)
    
mtx
```

Output:
```
[['a', 'pp', 'ppp', 'llll', 'eeeee'],
 ['p', 'll', 'uuu', 'mmmm'],
 ['p', 'ee', 'aaa', 'rrrr'],
 ['a', 'vv', 'ooo', 'cccc', 'aaaaa', 'dddddd', 'ooooooo']]
```

<!-- grade_id: cell-b2e0df6824c6b028, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
text = "ababaacdsadb"

char_freqs = {}

for c in text:
    try:
        char_freqs[c] += 1
    except KeyError:
        char_freqs[c] = 1
        
char_freqs
```

Output:
```
{'a': 5, 'b': 3, 'c': 1, 'd': 2, 's': 1}
```

```python
char_freqs2 = {}

for c in text:
    if c not in char_freqs2:
        char_freqs2[c] = 1
    else:
        char_freqs2[c] += 1

char_freqs2
```

Output:
```
{'a': 5, 'b': 3, 'c': 1, 'd': 2, 's': 1}
```

<!-- grade_id: cell-9af4ab90dccea74a, points: 0, grade: false, solution: true -->
```python
{c: text.count(c) for c in text}
```

Output:
```
{'a': 5, 'b': 3, 'c': 1, 'd': 2, 's': 1}
```

```python
d1 = {"a": 1, "b": 3, "c": 2}
d2 = {"a": 2, "b": 1}

d3 = {}

for key in set(d1.keys()) | set(d2.keys()):
    max_val = max(d1.get(key, 0), d2.get(key, 0))
    d3[key] = max_val

d3
```

<!-- grade_id: cell-269ea11b9d01f4dd, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## 2. Generators

### 2.1 Bigram generator
Write a generator function that takes a string and generates character bigrams from it:

<!-- grade_id: cell-e3f683ca2b1a38ff, points: 0, grade: false, solution: true -->
```python
def get_bigrams(text: str) -> str:
    for i, _ in enumerate(text[:-1]):
        yield text[i:i+2]
```

```python
gen = get_bigrams("abca")

for g in gen:
    print(g)
```

Output:
```
ab
bc
ca
```

<!-- grade_id: cell-a9c433466ba1aab5, points: 0, grade: true, solution: false -->
```python
assert list(get_bigrams("abca")) == ["ab", "bc", "ca"]
assert list(get_bigrams("")) == []
assert isinstance(get_bigrams("a"), GeneratorType)
```

### 2.2 Generators with memory

Write a generator function that remembers the previous element and yields it beside the current one.

<!-- grade_id: cell-993c1f55e59af7f8, points: 0, grade: false, solution: true -->
```python
def iter_with_memory(orig_iter):
    # YOUR CODE HERE
    raise NotImplementedError()
```

## ================ GRADED PART ================

## 3. Graded exercises

### 3.1 Formatted printing with comprehensions [6 points]

Given a list of lists of students, print them in the following format with one print invocation. You can find similar examples in Lecture 11.

```python
all_students = [
    ["Joe", "John", "Mary"],
    ["Tina", "Tony", "Jeff", "Bela"],
    ["Pete", "Dave"],
]
```

One class-per-line and print the size of the class too

~~~
class 1, size: 3, students: Joe, John, Mary
class 2, size: 4, students: Tina, Tony, Jeff, Bela
class 3, size: 2, students: Pete, Dave
~~~

<!-- grade_id: cell-cc231a351b5a051f, points: 0, grade: false, solution: true -->
```python
def print_students(all_students: List[List[str]]): 
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-d540370d23ed6e0a, points: 3, grade: true, solution: false -->
```python
f = io.StringIO()
with redirect_stdout(f):
    print_students(all_students)
s = f.getvalue()

assert s.strip() == "class 1, size: 3, students: Joe, John, Mary\nclass 2, size: 4, students: Tina, Tony, Jeff, Bela\nclass 3, size: 2, students: Pete, Dave"
```

Sort the classes by size in increasing order

~~~
class 3, size: 2, students: Pete, Dave
class 1, size: 3, students: Joe, John, Mary
class 2, size: 4, students: Tina, Tony, Jeff, Bela
~~~

<!-- grade_id: cell-7b6657c1ef2e4231, points: 0, grade: false, solution: true -->
```python
def print_ordered_students(all_students: List[List[str]]):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-cf7b5a337c36691a, points: 3, grade: true, solution: false -->
```python
f = io.StringIO()
with redirect_stdout(f):
    print_ordered_students(all_students)
s = f.getvalue()

assert s.strip() == "class 3, size: 2, students: Pete, Dave\nclass 1, size: 3, students: Joe, John, Mary\nclass 2, size: 4, students: Tina, Tony, Jeff, Bela"
```

### 3.2 Generators

#### 3.2.1 N-gram generator [3 points]

Change the function in the 2nd exercise so that it can generate n-grams. N should be a parameter of the function.

<!-- grade_id: cell-07f2eb9be62b41bf, points: 0, grade: false, solution: true -->
```python
def get_ngrams(text: str, n: int) -> str
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-4117b15c31f462b3, points: 3, grade: true, solution: false -->
```python
assert list(get_ngrams("abca", 2)) == ["ab", "bc", "ca"]
assert list(get_ngrams("abca", 3)) == ["abc", "bca"]
assert list(get_ngrams("abca", 3)) == ["abc", "bca"]
assert isinstance(get_ngrams("a", 1), GeneratorType)
```

#### 3.2.2 Generator with memory [6 points]
Write a generatior that takes an iterator and a memory size `k` and yields a tuple of k elements from the iterator.

<!-- grade_id: cell-213760247532dbf9, points: 0, grade: false, solution: true -->
```python
def iter_with_k_memory(orig_iter: Iterable, k: int) -> Tuple[Any]:
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-10d456c5035c4b76, points: 6, grade: true, solution: false -->
```python
l = list(range(3))
assert list(iter_with_k_memory(l, 3)) == [(0, None, None), (1, 0, None), (2, 1, 0)]
assert list(iter_with_k_memory(l, 1)) == [(0,), (1,), (2,)]
assert isinstance(iter_with_k_memory([], 1), GeneratorType)
```

### 3.3 Functional python

#### 3.3.1 Map [3 points]

Write a function that returns the uppercase version of a letter, if it was lower case and lower case if it was upper case. Use the map function to transform a string this way.

<!-- grade_id: cell-aef25ccd30b50c67, points: 0, grade: false, solution: true -->
```python
def upper_lower(character: str) -> str:
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-d67a26f3aaf855bf, points: 3, grade: true, solution: false -->
```python
assert "".join(list(map(upper_lower, "This iS A nOrmAl SentEnce"))) == "tHIS Is a NoRMaL sENTeNCE"
assert "".join(list(map(upper_lower, "AbSoLuTe MADness"))) == "aBsOlUtE madNESS"
```

#### 3.3.2 Filter [3 points]

Write a function that can be used with filter to filter out numeric characters in a string.

<!-- grade_id: cell-ec81b4f71ad41351, points: 0, grade: false, solution: true -->
```python
def just_alphabet(character: str) -> bool:
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-bec2c98a12c3ba20, points: 3, grade: true, solution: false -->
```python
assert "".join(list(filter(just_alphabet, "LE3T'S G0O"))) == "LET'S GO"
assert "".join(list(filter(just_alphabet, "4321"))) == ""
```

#### 3.3.3 Reduce [4 points]

Write a function that can concatenate any elements together as strings. Use the reduce function to demonstrate.

<!-- grade_id: cell-40d5eacebbe06cf7, points: 0, grade: false, solution: true -->
```python
def concat(string1: str, string2: Any) -> str:
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-d0e7ea0f483ba9d8, points: 4, grade: true, solution: false -->
```python
assert reduce(concat, ["I", "\'", "ve", " learned", " P", "ytho", "n", " <", 3], "") == "I've learned Python <3"
assert reduce(concat, [], "") == ""
assert reduce(concat, ["2", 2, 3.2, (3, 2), {2: 3}], "") == "223.2(3, 2){2: 3}"
```